Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Nhat Pham"
COLLABORATORS = ""
```

---

# CS110 Pre-class Work 5.1

## Question 1.

### Question 1a.

Write the worst sort function in the world, `worst_sort`. This function takes a list, and then randomly shuffles it until it happens to be in sorted order. Once it is in sorted order then the list is returned.


```python
def sort_check(A):
    for i in range(len(A)-1):
        if A[i]>A[i+1]:
            return True
            break
    return(False)
A = [1,2,3,2]
print(sort_check(A))
```

Output:
```
True
```

<!-- grade_id: cell-ead9c74054a1c5da, points: 0, grade: false, solution: true -->
```python
def worst_sort(A):
    """
    Sort A in ascending order by randomly shuffling its 
    elements until they are in order.
    
    Input:
    - A: list of numerical values
    
    Output:
    - A: Sorted list
    """
    import numpy as np
    while sort_check(A): # if the list is not sorted, shuffle the list
        np.random.shuffle(A)
    return(A)
#     raise NotImplementedError()
import random as rd
A = [rd.randint(0,9) for _ in range(9)]
worst_sort(A)
```

Output:
```
[5, 6, 7, 7, 7, 7, 8, 9, 9]
```

<!-- grade_id: cell-ff4db3a3e8b04515, points: 1, grade: true, solution: false -->
```python
# Please ignore this cell. This cell is for us to implement the tests 
# to see if your code works properly. 
```

### Question 1b.
What is the best case complexity of this algorithm?

<!-- grade_id: cell-afac80171b3b6232, points: 0, grade: true, solution: true -->
the best case is O(1), a constant. This case happpens when the input is sorted already

### Question 1c.

What is the average case complexity?


<!-- grade_id: cell-9d52d34daa1e3478, points: 0, grade: true, solution: true -->
The everage case complexity is O(n)*n! = O(n*n!)
O(n) is the running time of every shuffle as in sort_check function, we have to loop througgh every element to check if the list is sorted
There is only 1 sorted arrangement while n! permutations. Therefore, on average, we have to shuffle n! times to obtain that only sorted arragment, making the average times we have to run is n!. 

### Question 1d.

For what size lists is this a feasible method?

<!-- grade_id: cell-ab40f08768579798, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

### [Optional] Question 1e.

Can you think of an even worse sorting method? In such a case, what would its complexity be? How big a list could you sort?

<!-- grade_id: cell-ac05806942341926, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

## Question 2.

Approximate median finder. Given a list and δ (a number between 0 and 0.5), the approximate median finder function returns a number that is guaranteed to lie between the (50-tilde{δ}/2)% and (50+\tilde{δ}/2)% percentiles, where \tilde{δ}=100*δ. Implement such a function by randomly selecting an element from the list and testing whether or not it lies within the bounds. If it doesn’t then retry with a new random element, but only a limited number of retries are allowed (you decide the maximum number of retries.)

### Question 2a.
Complete the following function


<!-- grade_id: cell-2ab65512a6148d3c, points: 0, grade: false, solution: true -->
```python
def check_approx_med(A, num, delta):
    """
    Given a list and a number in a list, determine whether it is within (50-delta/2)% and 
    (50+delta/2)% percentiles.
    X% percentile of a list is defined to be the smallest number in the list that is as large as at least
    X% numbers in the list. 
    
    Input:
    - A: list of numerical values
    - num: the number of interest
    - delta: the error, lies between 0 and 0.5
    
    Output:
    - isApproxMed: a boolean value indicating whether num is within the bound
    """
    import numpy as np
    import random as rd
    #sort A to find percentile value
    A.sort()
    #determine percentiles with the given delta
    upper = (50 + delta*100/2)/100 
    lower = (50 - delta*100/2)/100
    #determine the index that satisfy the percentile requirements
    for i in reversed(range(1,len(A)+1)):
        if i/len(A) >= upper:
            upper_index = i-1
        if i/len(A) >= lower:
            lower_index = i-1
    #identify values corresponding to the above percentiles        
    upper = A[upper_index]
    lower = A[lower_index]
    #check
    if num >= lower and num <= upper:
        return True
    else:
        return False
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-df56d7433abb6517, points: 1, grade: true, solution: false -->
```python
assert(check_approx_med([0,1], 0, .25) == True)
assert(check_approx_med([0], 0, .25) == False)
```

Output:
```
AssertionError                            Traceback (most recent call last)
<ipython-input-32-aa1f56e94d4b> in <module>
      1 assert(check_approx_med([0,1], 0, .25) == True)
----> 2 assert(check_approx_med([0], 0, .25) == False)

AssertionError:
```

### Question 2b.
Complete the following function that makes use of `check_approx_med` above.


<!-- grade_id: cell-e19685b545147dde, points: 0, grade: false, solution: true -->
```python
def approx_med_finder(A, delta):
    """
    Given a list, find a number in the list that is between 50-delta/2% and 50+delta/2% percentiles
    
    Input:
    - A: list of numerical values
    - delta: the error, lies between 0 and 0.5
    
    Output:
    - num: the approximated median, if it is found within 100//delta trials (this is one possible 
    maximum number of retries. While you are encouraged to play around with another limits, the 
    submitted work must use this number.)
    - None: finding failed, if nothing found within 100//delta trials

    Note: the 100//delta is chosen here as it is the expected number of trials for the first successful finding to occur (using geometric distribution)
    """
    import numpy as np
    import random as rd
    for _ in range(100//delta):
        
        index = rd.randint(0,len(A)-1)
        num = A[index]
        upper = np.percentile(A,50 + delta*100/2)
        lower = np.percentile(A,50 - delta*100/2)
        if num >= lower and num <= upper:
            return num
        break
    return None
#     raise NotImplementedError()
```

### Question 2c.

What is the probability of failure in each random trial? What is the probability of failure after all the allowed random trials? Does this scale with δ or N (the number of elements in a list)?


<!-- grade_id: cell-efe7ab76f6a2546a, points: 0, grade: true, solution: true -->
prob of failure of one trial is the prob of getting a value outside the internval [(50-tilde{δ}/2)%, 50+\tilde{δ}/2)%], which is δ
The failure of the function happens when all trials fails, whose probability is δ^trials where trials is 100//delta. This result scales with δ

### Question 2d.
Analyze the expected runtime of `approx_med_finder`. Note that because the function uses `check_approx_med`, you will most likely need to analyze the runtime of that function, too.

<!-- grade_id: cell-86b199798477fc2a, points: 0, grade: true, solution: true -->
The running time of check_approx_med is O(n) as this function only requires a sort function and a constant number of operation to check in the input number is in within the range.
In approx_med_finder, we have at most 100//delta loops in which check_approx_med is called. This means the worst case running time is 100//delta*O(n) = O(100//delta*n) = O(n). The best case running time is O(n), which holds true when the first random number lies within the range.

```python

```