Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Nhat Pham"
COLLABORATORS = ""
```

---

# CS110 Pre-class Work 11.2

## Part A.  Currency trading (Slightly simplified version of 15.3-6 from Cormen et al.)

Imagine that you wish to exchange one currency for another. You realize that instead of directly exchanging one currency for another, you might be better off making a series of trades through other currencies, winding up with the currency you want.

Suppose that you can trade n different currencies, numbered $1,2,… ,n$, where you start with currency 1 and wish to wind up with currency $n$. You are given, for each pair of currencies $i$ and $j$ , an exchange rate $r_{ij}$ , meaning that if you start with $d$ units of currency $i$ , you can trade for $dr_{ij}$ units of currency $j$. Note that the total number of trades allowed is limited to $n$.

Assuming there is no commission, propose an algorithm (described in prose) to solve this problem using either a memoization or bottom-up strategy. **Note that we will be working on the Python implementation in class.**

Click [here](https://drive.google.com/open?id=1L8Fjo1Xt8sltab-tz3m91eTuiE9LYWF7) for some example data.

<!-- grade_id: cell-3f8a8808c720baf2, points: 0, grade: true, solution: true -->
the currency rate table can translate to a network of nodes (currency) and edges (whose weight is the direct exchange rate between 2 node). So the problem is equivalent to finding a simple path of which the product of the weights of edges is largest. The simple path (a path passing a node only once) is included as an constraint because there is no point buying a currency then selling it for the currency we used to have.To store results computed during the bottom up process, a matrix n*n is created where c(i,j) record the largest currency to buy j using i. We number the currency we want to buy n and sell 1. For example, there are 5 provided currencies, and we want to buy Pound using USD, then USD is 1 and pound is 5. Other currencies take to any number between 1 and 5. Accordingly a matrix 5*5 is created to store the max exchange rate between 5 currencies

* To avoid repeating subproblems, we need to fill in the max exchange rate in cells of which the difference of i and j is 1, then 2, 3, ,4
* when |i-j| = 1, the largest exchange rates are the provided ones in the table as we have no intermediate node to insert between node i and j
* When |i-j| = 2 (i.e i = 1, j = 3), we need to compare path 1-3 and (1-2)-(2-3). the max of these two is filled in c(1,3)
* When |i-j| = 3 (i.e i = 1, j = 4), we need to compare more paths including 1-4,1-2-2-4, 1-3-3-4. We don need to consider path 1-2-3-4 as its result is included in 1-2-2-4 or 1-3-3-4 already. Because if 1-2-3 is better than 1-3, the product of 1-2 and 2-3 is filled in c(1,3) not the weight of edge 1-3.
* When |i-j| = 4, (i=1 j = 5). This is the most important cell as it tells how much at most we can buy Pound using USD. The process is similar to smaller values of |i-j|. We compute the exchange rate of all intermediate paths and get the max out of such results to fill in c(i,j)

## Part B - Money game (Solution)
Consider a row of n coins of values $v_1, v_2,...,v_n$, **where $n$ is even**. We play a game against an opponent by alternating turns. In each turn, a player selects either the first or last coin from the row, removes it from the row permanently, and receives the value of the coin. Determine the maximum possible amount of money we can definitely win if we move first.

For example, consider the game:

\$2, \$10, \$1, \$5

By moving first and playing optimally one can be guaranteed of \$15. The first move is to take \$5. This forces your opponent to take either \$2 or \$1, and then allows you to take \$10. Assume that the opposing player also plays optimally (i.e., minimizing the gain of the first-move player.)

As a hint, we will provide the recurrence for the solution to this problem. Suppose $c(i,j)$ is the maximum possible amount of money the first player can win for the row of coins of value $v_i, v){i+1},...,v_j$ (so to solve our problem, we need to compute $c(1,n)$. Then:

$$c(i,j)=max[v_i+min(c(i+2,j), c(i+1,j-1)), v_j+min(c(i,j-2), c(i+1,j-1))]$$

You are also encouraged to watch [this video](https://www.youtube.com/watch?v=KnP8_L13xW4&list=PLF_a-qBXTGFektoI6JUOTRL36JlvD04BR&index=5) if you need some more hints.

## Question 1. 
Complete the following function to solve the game using a bottom-up strategy, assuming that the opposing player also plays optimally. 


<!-- grade_id: cell-4014ae901b66304b, points: 0, grade: false, solution: true -->
```python
def bottom_up_coin_game(A):
    """
    Returns the maximum possible amount of money the first-move player can win,
    given an array of coin values. The function solves this using a bottom-up 
    approach.
    
    Inputs:
    - A: list of floats, values of the coins, of even length. 
    
    Outputs:
    - max_val: float, maximum possible amount of money that can be won by the first 
    player. max_val is None when the length of list A is odd.
    """
    # YOUR CODE HERE
    n = len(A) #len of the list
    if n%2 == 1: # if the length is not even
        return None
    
    if n == 0: #if len of is 0
        return 0
    
    store = [[0 for _ in range(n)] for _ in range(n)] # creat a matrix to store computed result
    for delta in range(n): #difference of the start and the end of the sublist (input list of subproblem)
        for i in range(n-delta): #start of the sublist
            j = i + delta  #end of the sublist
            if delta <= 1: #if len of sublist is no greater than 1
                store[i][j]= max(A[i:j+1])
#                 print(i,j,store) #check the process
            if delta > 1: #else
                store[i][j] = max(A[i]+min(store[i+2][j], store[i+1][j-1]), 
                                  A[j]+min(store[i][j-2], store[i+1][j-1]))
#                 apply the approach in the video, 
#                 note that results of subproblems now are stored in the array store
#                 print(i,j,store) #check the process
    return store[0][n-1]
bottom_up_coin_game([2, 10, 1, 5])
#     return maxsum
#         raise NotImplementedError()
```

Output:
```
15
```

```python
def coin_game(A):
    """
    Returns the maximum possible amount of money the first-move player can win,
    given an array of coin values. The function solves this using a bottom-up 
    approach.
    
    Inputs:
    - A: list of floats, values of the coins, of even length. 
    
    Outputs:
    - max_val: float, maximum possible amount of money that can be won by the first 
    player. max_val is None when the length of list A is odd.
    """
    # YOUR CODE HERE
    if len(A)%2 == 1:
        return None
    if len(A) == 0:
        return(0)
    if len(A) == 1:
        return (A[0])
    if len(A) == 2:
        return (max(A[0],A[1]))
    return(max(A[0]+min(coin_game(A[2:]), coin_game(A[1:-1])),
                     A[-1]+min(coin_game(A[0:-2]), coin_game(A[1:-1]))))

#     return maxsum
#         raise NotImplementedError()
```

<!-- grade_id: cell-f7a5d6c26e95dab2, points: 1, grade: true, solution: false -->
```python
import numpy as np
assert(bottom_up_coin_game([2, 10, 1, 5]) == 15)  
assert(bottom_up_coin_game([2, 10, 1, 5, 5]) is None)  
assert(bottom_up_coin_game([]) == 0)  
```

## [Optional]  Question 2. 

Complete the function `print_strategy` to print out a solution to the coin game. Completing the helper function `extended_bottom_up_coin_game` to use in `print_strategy` is optional.

<!-- grade_id: cell-4583755bc245e1e9, points: 0, grade: false, solution: true -->
```python
def print_strategy(A):
    """
    Print coins to pick for the first player
    
    Inputs:
    - A: list of floats, values of the coins, of even length.
    
    Outputs:
    - max_val: float, maximum possible amount of money that can be won by the first 
    player. 
        * max_val is None if the length of list A is odd.
        * max_val is 0 if list A is empty (no coins)
    - out: list of values of the coins that the first player picks in that
    order.
        * out is None if max_val is 0 or None.
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-75ce551c5185e374, points: 0, grade: true, solution: false -->
```python
assert(print_strategy([4,6,5,20,4,6,4,8])[0] == 40)
assert(print_strategy([4,6,5,20,4,6,4,8])[1] == [8, 6, 6, 20])
```