Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Nhat Pham"
COLLABORATORS = ""
```

---

# CS110 Pre-class Work 3.1

## Question 1. 
Paste in your Python implementation of the maximum subarray from the previous class in the cell below and use that to find out the value of the maximum subarray of this array: `A = [-2, -3, 4, -1, -2, 1]`

```python
def bruteforce_max_subarray(A):
    max = float('-inf')
    profit =()
    for i in range(0,len(A)-1):
        for j in range(i+1,len(A)):
            if sum(A[i:j+1]) > max:
                max = sum(A[i:j+1])
                profit = (i,j,max)
    return(profit)
                
A = [-2, -3, 4, -1, -2, 1]
bruteforce_max_subarray(A)

```

Output:
```
(2, 3, 3)
```

<!-- grade_id: cell-e69c8bbcd40ca242, points: 0, grade: true, solution: true -->
```python
def max_cross_subarray(A,low,mid,high):
    left_max = right_max = float('-inf')
    sum = 0
    for i in reversed(range(mid)):
        sum += A[i]
        if sum > left_max:
            left_max = sum
            max_left = i
    sum = 0
    for j in range(mid,high,1):
        sum += A[j]
        if sum > right_max:
            right_max = sum
            max_right = j
    return(max_left,max_right,left_max+right_max)
def max_subarray(A,low,high):
    if low == high:
        return(low,high,A[low])
    if low != high:
        mid = (low+high)//2
    left_low,left_high,left_max = max_subarray(A,low,mid)
    right_low,right_high,right_max = max_subarray(A,mid+1,high)
    cross_low,cross_high,cross_max = max_cross_subarray(A,low,mid,high)
    if leff_max == max(left_max,right_max,cross_max):
        return(left_low,left_high,left_max)
    elif right_max == max(left_max,right_max,cross_max):
        return(right_low,right_high,right_max)
    elif cross_maxx == max(left_max,right_max,cross_max):
        return(cross_low,cross_high,cross_max)
# raise NotImplementedError()
```

## Question 2. 
Now, your friend Joe comes and appends a single extra number at the end of the array, which becomes:  `B = [-2, -3, 4, -1, -2, 1, 8]`. Do you have to re-run the entire maximum subarray again? Explain your answer. 
The subsequent questions will help you figure out an efficient algorithmic strategy to address the last question, but make sure to write your explanation above first, before answering the remaining questions.


<!-- grade_id: cell-d4b7cd845c816263, points: 0, grade: true, solution: true -->
Given the brute force method above,we have to re-run all over again because we increase the number of inner and outer loops by one each, meaning more combinations to consider.


## Question 3. 

**Determine if the following statement is True or False and explain your answer.**
If the maximum subarray of the array A is different than the maximum subarray of the array B (questions 1 and 2), the new maximum-subarray doesn’t need to include 8 (i.e., the newly appended element). 


<!-- grade_id: cell-9377964a8756f13b, points: 0, grade: true, solution: true -->
it has to include 8 because there is no reason for the different return otherwise. To be specific, in the old result, we already determine the best buying day. Given a new element of the list, we just need to consider the old max value with the sum of the list beggining from the pre-determined buying day to the newl-added element.

## Question 4.
Complete the Python function `incremental_max_subarray(x, mx)` in the cell below.

This [video](https://www.youtube.com/watch?v=AAgErqQmwmA&list=PLF_a-qBXTGFektoI6JUOTRL36JlvD04BR&index=4&t=0s)  might be helpful to understand the `incremental_max_subarray` problem.

<!-- grade_id: cell-0230e459f3d701f8, points: 0, grade: false, solution: true -->
```python
def incremental_max_subarray(x,mx):
    if len(x)==1:
        return(x[0])
    else:
        max_current=[x]
        for i in range(len(x)):
            if sum(x[i:]) >= mx:
                mx = (sum(x[i:]))
    return(mx)

                   
    """ 

    Inputs:
    - x: a NON-EMPTY list of numbers (e.g., the array B in the first two questions above). 
        * If x has 1 element: returns the value of the element regardless of the value of mx
    - mx: the maximum subarray of x excluding its last element (i.e., compute the 
    maximum subarray of the input array x considering only its first len(x) - 1 elements)
    
    Output: The maximum subarray of the array x.
    
    For example, using the array B in question 2, the result of incremental_max_subarray(B, 4) 
    is 10 (10 = 8 + 1 - 2 -1 + 4). 
    """
B = [-2, -3, 4, -1, -2, 1, 8]
incremental_max_subarray(B,4)
    # YOUR CODE HERE
#     raise NotImplementedError()
```

Output:
```
10
```

<!-- grade_id: cell-9abf19e56620ffa3, points: 1, grade: true, solution: false -->
```python
B = [-2, -3, 4, -1, -2, 1, 8]
assert(incremental_max_subarray(B, 4) == 10)
assert(incremental_max_subarray(B[:1], 0) == B[0])
```

<!-- grade_id: cell-bacebe71191cbd0f, points: 0, grade: true, solution: true -->
```python

```

<!-- grade_id: cell-fd96d4ccccd99fe6, points: 0, grade: false, solution: true -->
```python
def max_subarray(A):
    mx = float('-inf')
    global_max = float('-inf')
    for j in range(1,len(A)+1):
        print('A[i:j]',A[i:j])
        print('mx=',mx)
        print('incremental_max_subarray(A[i:j],mx)',incremental_max_subarray(A[i:j],mx))
        mx = incremental_max_subarray(A[i:j],mx)
        if mx > global_max:
            global_max = mx
    return(global_max)
            
            
    """
    Using `incremental_max_subarray` iteratively on A to produce the value of the maximum
    subarray of A.
    
    Inputs:
    - A: a NON-EMPTY list of floats
    
    Outputs: flaot, the sum of the maximum subarray of A
    """
    # YOUR CODE HERE
#     raise NotImplementedError()
max_subarray(B)
```

Output:
```
A[i:j] []
mx= -inf
incremental_max_subarray(A[i:j],mx) -inf
A[i:j] [-3]
mx= -inf
incremental_max_subarray(A[i:j],mx) -3
A[i:j] [-3, 4]
mx= -3
incremental_max_subarray(A[i:j],mx) 4
A[i:j] [-3, 4, -1]
mx= 4
incremental_max_subarray(A[i:j],mx) 4
A[i:j] [-3, 4, -1, -2]
mx= 4
incremental_max_subarray(A[i:j],mx) 4
A[i:j] [-3, 4, -1, -2, 1]
mx= 4
incremental_max_subarray(A[i:j],mx) 4
A[i:j] [-3, 4, -1, -2, 1, 8]
mx= 4
incremental_max_subarray(A[i:j],mx) 10
10
```

```python
[1,3,5,3,6,2].index(3)
```

Output:
```
1
```

<!-- grade_id: cell-669ad779766aa2c3, points: 1, grade: true, solution: false -->
```python
# Please ignore this cell. This cell is for us to implement the tests 
# to see if your code works properly. 
```

### Part 2.
Is this more efficient than the divide-and-conquer approach? Explain.

<!-- grade_id: cell-cd8f0b130a7136db, points: 0, grade: true, solution: true -->
in this method, we run 2 loops, for each of which the running time depends on its input size. this means the run time is \theta(n^2) while that of the divide-and-conquer method is \theta(nlog(n)). Therefore, the dividie-and-conquer method is more efficient.

```python

```