Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Nhat Pham"
COLLABORATORS = ""
```

---

# CS110 Pre-class Work 13.1

## Question 1.

Complete the following functions, following the algorithms in Cormen et al.

<!-- grade_id: cell-542b2dce56293aef, points: 0, grade: false, solution: true -->
```python
def lcs_length(x, y):
    """
    Computes the length of an LCS of strings x and y.
    
    Inputs:
    - x, y: strings
    
    Outputs:
    - c: a list of lists of ints OR a numpy array. c[i,j] contains the 
    length of a LCS of x[:i] and y[:j]
    - b: a list of lists of strings OR a numpy array, containing the information
    used for LCS reconstruction (See Cormen et al.) Use "N" (North), "NW" 
    (North West), and "W" (West) that correspond to the directions of the arrows 
    used in Cormen et al.
    """
    m = len(list(x))
    n = len(list(y))
    #a list of sublists to store the length of lcs
    b = [[None for _ in range(n)] for _ in range(m)]
    #a list of sublists to store the direction
    c = [[None for _ in range(n+1)] for _ in range(m+1)]
    # fill in row 0 and col 0 with 0
    for i in range(1,m+1):
        c[i][0] = 0
    for j in range(n+1):
        c[0][j] = 0
    #fill in the list the lsc of subproblems
    for i in range(1,m+1):
        for j in range(1,n+1):
            # if the last characters of 2 sublist match
            if list(x)[i-1] == list(y)[j-1]:
                c[i][j] = c[i-1][j-1] + 1
                #record in b table where the result of c[i][j] comes from
                b[i-1][j-1] = 'NW' 
            #if last characters of 2 sublist do not match, take max(c[i-1][j],c[i][j-1])) to fill in c[i][j]
            elif c[i-1][j] >= c[i][j-1]:
#                 print(i,j)
                c[i][j] = c[i-1][j]
                b[i-1][j-1] = 'N'
            else:
                c[i][j] = c[i][j-1]
                b[i-1][j-1] = 'W'
    return c,b
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
lcs_length('ambgdec', 'aubyci')
```

Output:
```
([[0, 0, 0, 0, 0, 0, 0],
  [0, 1, 1, 1, 1, 1, 1],
  [0, 1, 1, 1, 1, 1, 1],
  [0, 1, 1, 2, 2, 2, 2],
  [0, 1, 1, 2, 2, 2, 2],
  [0, 1, 1, 2, 2, 2, 2],
  [0, 1, 1, 2, 2, 2, 2],
  [0, 1, 1, 2, 2, 3, 3]],
 [['NW', 'W', 'W', 'W', 'W', 'W'],
  ['N', 'N', 'N', 'N', 'N', 'N'],
  ['N', 'N', 'NW', 'W', 'W', 'W'],
  ['N', 'N', 'N', 'N', 'N', 'N'],
  ['N', 'N', 'N', 'N', 'N', 'N'],
  ['N', 'N', 'N', 'N', 'N', 'N'],
  ['N', 'N', 'N', 'N', 'NW', 'W']])
```

<!-- grade_id: cell-6b825bf900c2cf0b, points: 0, grade: false, solution: true -->
```python
def print_lcs(b,x,i,j):
    """
    Finds a LCS.
    
    Inputs:
    - b: a list of lists of strings OR a numpy array, returned by lcs_length
    - x: string, an input to lcs_length
    - i, j: ints. print_lcs(b,x,i,j) returns a lcs of x[:i] and y[:j], where y
    is an input to lcs_length.
    
    Outputs:
    - lcs: list of strings, representing a LCS of x and y
    - length: int, the length of the LCS
    
    You can choose to actually PRINT OUT the LCS or not using the print function.
    
    """
    #the base case is when points are to the left of the first letter of the string (i,j = -1)
    if i<0 or j<0:
        return
    # b[i][j] == 'NW' means the subproblem has a match at the end of 2 substrings
    #-> return the common letter at the end of the first input string (X)
    if b[i][j] == 'NW':
        print(i,j)
        print_lcs(b,x,i-1,j-1)
        print(list(x)[i])
    #if there is no common letter at the end of two substrings, 
    #remove the last character of one substring to create a new subproblem depending on the value in b table
    elif b[i][j] == 'N':
        print(i,j)
        print_lcs(b,x,i-1,j)
    else:
        print(i,j)
        print_lcs(b,x,i,j-1)
#     return(letter)
    
#     raise NotImplementedError()
```

```python
lcs_length(x, y)
```

Output:
```
([[0, 0, 0, 0, 0, 0, 0],
  [0, 1, 1, 1, 1, 1, 1],
  [0, 1, 1, 1, 1, 1, 1],
  [0, 1, 1, 2, 2, 2, 2],
  [0, 1, 1, 2, 2, 2, 2],
  [0, 1, 1, 2, 2, 2, 2],
  [0, 1, 1, 2, 2, 2, 2],
  [0, 1, 1, 2, 2, 3, 3]],
 [['NW', 'W', 'W', 'W', 'W', 'W'],
  ['N', 'N', 'N', 'N', 'N', 'N'],
  ['N', 'N', 'NW', 'W', 'W', 'W'],
  ['N', 'N', 'N', 'N', 'N', 'N'],
  ['N', 'N', 'N', 'N', 'N', 'N'],
  ['N', 'N', 'N', 'N', 'N', 'N'],
  ['N', 'N', 'N', 'N', 'NW', 'W']])
```

```python
def sub_print_lcs(b,x,i,j,result):
    """
    Finds a LCS.
    
    Inputs:
    - b: a list of lists of strings OR a numpy array, returned by lcs_length
    - x: string, an input to lcs_length
    - i, j: ints. print_lcs(b,x,i,j) returns a lcs of x[:i] and y[:j], where y
    is an input to lcs_length.
    
    Outputs:
    - lcs: list of strings, representing a LCS of x and y
    - length: int, the length of the LCS
    
    You can choose to actually PRINT OUT the LCS or not using the print function.
    
    """
    if i<0 or j<0:
        return 
    if b[i][j] == 'NW':
#         print(i,j)
#         print('here',list(x)[i] )
#         result = [list(x)[i]] + result
        sub_print_lcs(b,x,i-1,j-1,result)
        result.append(list(x)[i])
#         result.append(list(x)[i-1])
    elif b[i][j] == 'N':
#         print(i,j)
        sub_print_lcs(b,x,i-1,j,result)
    else:
#         print(i,j)
        sub_print_lcs(b,x,i,j-1,result)
    return(result)

#let sub_print_lcs(b,x,i,j,result) be the recrusive function so that a list, named result, is created
#inside the main function, named print_lcs(b,x,i,j), to store required outputs

def print_lcs(b,x,i,j):
    result = []
    sub_print_lcs(b,x,i,j,result)
    result= [result,len(result)]
    return result
#     raise NotImplementedError()

###
import numpy as np
x, y = 'ambgdec', 'aubyci'
c, b = lcs_length(x, y)
print_lcs(b,x,len(x)-1,len(y)-1)
```

Output:
```
[['a', 'b', 'c'], 3]
```

<!-- grade_id: cell-e06e8f09bb8a5063, points: 1, grade: true, solution: false -->
```python
import numpy as np
x, y = 'ambgdec', 'aubyci'
c, b = lcs_length(x, y)
assert(print_lcs(b,x,len(x)-1,len(y)-1)[0] == ['a', 'b', 'c'])
assert(print_lcs(b,x,len(x)-1,len(y)-1)[1] == 3)

x, y = 'xyqwsssazdesaqqf', 'xoppoypllzookjdef'
c, b = lcs_length(x, y)
assert(print_lcs(b,x,len(x)-1,len(y)-1)[0]  == ['x', 'y', 'z', 'd', 'e', 'f'])
assert(print_lcs(b,x,len(x)-1,len(y)-1)[1]  == 6)

```

## Question 2. (Adapted from Exercise 15-4.1 Cormen et al.)
Use the functions built in Question 1 to find the LCS of `'10010101'` and `'010110110'`. You should store the list that represents the LCS you found in a variable named `lcs_q2`

<!-- grade_id: cell-5d742a82dde1229f, points: 0, grade: true, solution: true -->
```python
def sub_print_lcs(b,x,i,j,result):
    """
    Finds a LCS.
    
    Inputs:
    - b: a list of lists of strings OR a numpy array, returned by lcs_length
    - x: string, an input to lcs_length
    - i, j: ints. print_lcs(b,x,i,j) returns a lcs of x[:i] and y[:j], where y
    is an input to lcs_length.
    
    Outputs:
    - lcs: list of strings, representing a LCS of x and y
    - length: int, the length of the LCS
    
    You can choose to actually PRINT OUT the LCS or not using the print function.
    
    """
    if i<0 or j<0:
        return 
    if b[i][j] == 'NW':
#         print(i,j)
#         print('here',list(x)[i] )
#         result = [list(x)[i]] + result
        sub_print_lcs(b,x,i-1,j-1,result)
        result.append(list(x)[i])
#         result.append(list(x)[i-1])
    elif b[i][j] == 'N':
#         print(i,j)
        sub_print_lcs(b,x,i-1,j,result)
    else:
#         print(i,j)
        sub_print_lcs(b,x,i,j-1,result)
    return(result)


def print_lcs(b,x,i,j):
    lcs_q2 = []
    sub_print_lcs(b,x,i,j,lcs_q2)
    lcs_q2= [lcs_q2,len(lcs_q2)]
    return lcs_q2
# raise NotImplementedError()

```

```python
x, y = '51042679', '01245679'
c, b = lcs_length(x, y)
print_lcs(b,x,len(x)-1,len(y)-1)
```

Output:
```
[['1', '4', '6', '7', '9'], 5]
```

## Question 3. (Adapted from Exercise 15-4.5 Cormen et al.)
Complete the following function, making use of `lcs_length` and `print_lcs`. 



The problem is rephrased as finding the the lcs of the input list and its arranged version (lst vs ascending lst)
And we use three_way_merge_sort algorithm to sort the list and assigned the result to a new variable called arranged_lst.
This approach allows us to find the lmis of a list with O(n^2) because we reduce the length of 2 n-long inputs by one unit each to create n^2 subproblems to find their lcs. As a result, the complexity of this result scales with O(n^2) where n is the length of the input

```python
def merge(list1,list2,list3):
#     Change the type of all inputs to list so that the code can work with array
    L = list(list1)
    M = list(list2)
    R = list(list3)
#     Append sentinel values to make sure the below comparisons always work
    L.append(float('inf'))
    M.append(float('inf'))
    R.append(float('inf'))

    i = j = k = 0
#   Determine the lenth of the sorted list, which is the sum of the length of 3 sublists. 
#   -3 is used to not add 3 inf values above, which should not be in the sorted list
    length = len(L)+len(M)+len(R)-3
    A = []
    
#   loop m through the range [0:length -1] to make sure the sorted list has all elements of sublists
    for m in range(length):
#       find the min value in 3 values in 3 sublists and append it to the sorted list A
        if L[i] == min(L[i],M[j],R[k]):
            A.append(L[i])
            i += 1
        elif M[j] == min(L[i],M[j],R[k]):
            A.append(M[j])
            j += 1
        else:
            A.append(R[k])
            k += 1
    return(A)

def three_way_merge(A):
# calculate index to split the input in three parts
    r = len(A)
    mid_1 = r//3 # 
    mid_2 = 2*r//3 # 
    if len(A) > 1:
#       Recall the function to brake the list to the base case of 1 element
        list1 = three_way_merge(A[:mid_1])
        list2 = three_way_merge(A[mid_1:mid_2]) 
        list3 = three_way_merge(A[mid_2:r])

#       recal the merge funcntion to merge sort sublists
        output = merge(list1, list2, list3)

    else:
        output = A
    return(output)

```

<!-- grade_id: cell-429c944f34a54c91, points: 0, grade: true, solution: true -->
```python

def lmis(lst):
    arranged_lst = three_way_merge(lst)
    
    x, y = lst, arranged_lst
    c, b = lcs_length(x, y)
    LMIS = print_lcs(b,x,len(x)-1,len(y)-1)
    return(list(LMIS[0]))
    """
    Finds the Longest Monotonically Increasing Subsequence (LMIS) of a list 
    (lst) of n numbers in O(n^2) time. Note that a monotonically increasing 
    sequence is a sequence of numbers such that: a_1 <= a_2 <= ... <= a_n .
    
    Inputs:
    - lst: a list of ints
    
    Outputs:
    - out_lst: a list of ints, a longest monotonically increasing subsequence
    of lst
    """
    # YOUR CODE HERE
#     raise NotImplementedError()
lmis([5, 1, 0, 4, 2, 6, 7, 9])
```

Output:
```
[1, 4, 6, 7, 9]
```

<!-- grade_id: cell-f8b218898fb43d19, points: 0, grade: true, solution: false -->
```python
assert(lmis([5, 1, 0, 4, 2, 6, 7, 9]) == [1, 4, 6, 7, 9])
assert(lmis([6, -1,4,5,5,7,-6,2]) == [-1, 4, 5, 5, 7])
```

```python

```