Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Nhat Pham"
COLLABORATORS = ""
```

---

# CS110 Pre-class Work 3.2

## Question 1.
Given the array `H=[39, 85, 85, 16, 49, 7, 49, 92, 76, 15, 21, 30, 29, 31, 28]`, perform the following operations:
1. Draw the corresponding binary tree of H. Is the binary tree a valid max heap? Explain your answer.
2. Using as a model the drawing examples illustrated in Figure 6.2 of Cormen et al.,  draw a step-by-step transformation of the array above into a valid max heap. 
3. Now that you have obtained a valid max heap, write out the corresponding array that stores the valid max-heap.

Use as many cells as you wish for this question.

<!-- grade_id: cell-b1f821b034b44619, points: 0, grade: true, solution: true -->
* refer to this link :https://docs.google.com/document/d/1NiqOW068gXWsEdtMWMqx-MlhR0oZ5EyaEPPsLaZGDPU/edit?usp=sharing
* new list: [92, 85, 85, 76, 49, 30, 49, 16, 39, 15, 21, 7, 29, 31, 28]

## Question 2. 
Consider the following questions on the $MAX-HEAPIFY$ operation.
### Question 2a.

In the pseudocode of $MAX-HEAPIFY$ (Cormen et al., p.154, or you can view it [here](https://drive.google.com/open?id=1e_3jsX4-qQCfZXKMok_T6LPFh9FwtmT5)), what does A.heap-size mean and what is the idea behind the local variable largest? 


<!-- grade_id: cell-06106b81a909b9e6, points: 0, grade: true, solution: true -->
Heap_size is the number of elements being sorted by the max-heap method. The difference  is sometimes heap-size is smaller than len(A) as we do need to sort all the list at the beginning, 
The largest variable is the index of parent which is supposed to be larger than any of its children.

### Question 2b.
The functions $LEFT(i)$ and $RIGHT(i)$, lines 1 and 2 in the $MAX-HEAPIFY$ pseudocode, return the array index of the left and right child, respectively, of a node in a binary tree. From reading Section 6.1, you know that the input to both functions is an integer number, $i$, which corresponds to the array index of the parent node in the array. Review Section 6.1 for more information. Write a Python implementation of the functions $LEFT(i)$ and $RIGHT(i)$ by filling in the cells below.

<!-- grade_id: cell-2efd4017321c8d48, points: 0, grade: false, solution: true -->
```python
def left(i):
    return (2*(i+1)-1)
    raise NotImplementedError()
```

<!-- grade_id: cell-507ddbab62e6a32c, points: 1, grade: true, solution: false -->
```python
# Please ignore this cell. This cell is for us to implement the tests 
# to see if your code works properly. 
```

<!-- grade_id: cell-b7e5ceeacaeaf1b8, points: 0, grade: false, solution: true -->
```python
def right(i):
    return (2*(i+1))
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-c24867515653788a, points: 1, grade: true, solution: false -->
```python
# Please ignore this cell. This cell is for us to implement the tests 
# to see if your code works properly. 
```

### Question 2c.
Write a Python implementation of the MAX-HEAPIFY operation using the pseudocode above, and your newly written functions, `left` and `right`.

<!-- grade_id: cell-52289253a243341b, points: 0, grade: false, solution: true -->
```python
def heapify(A, i):
    heap_sz = len(A)
    l = left(i)
    r = right(i)
#     print(heap_sz,l,r)
    if l+1 <= heap_sz and A[l] > A[i]:
        largest = l
    else:
        largest = i
#     print(r, largest)
    if r+1 <= heap_sz and A[r] > A[largest]:
        largest = r
    if largest != i:
        A[i], A[largest] = A[largest], A[i]
        heapify(A,largest)
    return(A)
    """
    Inputs:
    - heap: a list of floats. Assume that the heap size is the length of the heap.
    
    No output is needed. This function should modify (if necessary) heap in-place.
    """
    # YOUR CODE HERE
#     raise NotImplementedError()
A = [39, 85, 85, 16, 49, 7, 49, 92, 76, 15, 21, 30, 29, 31, 28]
heapify(A,0)

```

Output:
```
[85, 49, 85, 16, 39, 7, 49, 92, 76, 15, 21, 30, 29, 31, 28]
```

<!-- grade_id: cell-7f63745ff34b881b, points: 1, grade: true, solution: false -->
```python
A = [39, 85, 85, 16, 49, 7, 49, 92, 76, 15, 21, 30, 29, 31, 28]
heapify(A,0)
assert(A == [85, 49, 85, 16, 39, 7, 49, 92, 76, 15, 21, 30, 29, 31, 28])

```

## Question 3. 
Next, write a Python implementation of the BUILD_MAX_HEAP operation using the pseudocode provided in Section 6.3 of Cormen et. al. Test your Python implementation using the array in problem 1, and make sure your Python codes produce a valid max heap.

<!-- grade_id: cell-9237da6268a5e3eb, points: 0, grade: false, solution: true -->
```python
def build_max_heap(A):
    heap_sz=len(A)
    for i in range((len(A)-1)//2,-1,-1):
        heapify(A,i)
    """
    Input:
    - A: a list of floats.
    
    No output is needed. The function should turn A into a valid max heap, in-place.
    """
A = [4, 1, 3, 2, 16, 9, 10, 14, 8, 7]
build_max_heap(A)
print(A)
H=[39, 85, 85, 16, 49, 7, 49, 92, 76, 15, 21, 30, 29, 31, 28]
build_max_heap(H)
print(H)

    # YOUR CODE HERE
#     raise NotImplementedError()
```

Output:
```
[16, 14, 10, 8, 7, 9, 3, 2, 4, 1]
[92, 85, 85, 76, 49, 30, 49, 16, 39, 15, 21, 7, 29, 31, 28]
```

<!-- grade_id: cell-5dc147b837da3f7e, points: 1, grade: true, solution: false -->
```python
A = [4, 1, 3, 2, 16, 9, 10, 14, 8, 7]
build_max_heap(A)
assert(A == [16, 14, 10, 8, 7, 9, 3, 2, 4, 1])
```

## Question 4. 

Lastly, write Python implementations of the $MIN-HEAPIFY$ and $BUILD-MIN-HEAP$ operations for a min heap data structure. You can use your $MAX-HEAPIFY$ and $BUILD-MAX-HEAP$ Python function as models, just remember that the latter two functions support operations for the max heap data structure. Test your Python implementation of $BUILD-MIN-HEAP$ using the array in problem 1, and make sure your Python codes produce a valid min heap. 

<!-- grade_id: cell-d86943f127876fc8, points: 0, grade: false, solution: true -->
```python
def min_heapify(A, i):
    heap_sz = len(A)
    l = left(i)
    r = right(i)
#     print(heap_sz,l,r)
    if l+1 <= heap_sz and A[l] < A[i]:
        smallest = l
    else:
        smallest = i
#     print(r, smallest)
    if r+1 <= heap_sz and A[r] < A[smallest]:
        smallest = r
    if smallest != i:
        A[i], A[smallest] = A[smallest], A[i]
        min_heapify(A,smallest)
    return(A)
A = [39, 85, 85, 16, 49, 7, 49, 92, 76, 15, 21, 30, 29, 31, 28]
min_heapify(A,0)
    
    # YOUR CODE HERE
#     raise NotImplementedError()
```

Output:
```
[39, 85, 85, 16, 49, 7, 49, 92, 76, 15, 21, 30, 29, 31, 28]
```

<!-- grade_id: cell-6efa45e3956ad0b8, points: 0, grade: false, solution: true -->
```python
def build_min_heap(A):
    heap_sz = len(A)
    for i in range((len(A)-1)//2,-1,-1):
        min_heapify(A,i)

A = [39, 85, 85, 16, 49, 7, 49, 92, 76, 15, 21, 30, 29, 31, 28]
build_min_heap(A)
print(A)


    # YOUR CODE HERE
#     raise NotImplementedError()
```

Output:
```
[7, 15, 28, 16, 21, 29, 31, 92, 76, 49, 85, 30, 85, 39, 49]
```

<!-- grade_id: cell-7e64941402f819c2, points: 1, grade: true, solution: false -->
```python
# Please ignore this cell. This cell is for us to implement the tests 
# to see if your code works properly. 
```