Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Nhat Pham"
COLLABORATORS = ""
```

---

# CS110 Pre-class Work 7.1

## Part A. Median Heap (watch this [video explanation](https://www.youtube.com/watch?v=756_8C9YBZQ&list=PLF_a-qBXTGFektoI6JUOTRL36JlvD04BR&index=5&t=0s) or read this [description](https://stackoverflow.com/a/15319593/7946759))

Throughout this pre-class work, please use the following definition of median: the median of a list of numbers is the one in the middle of the list when the list is ordered. When such the middle element can’t be determined (i.e., in a list of even length), the average of the two middle elements is the median. For example, 5 is the median of [-1,2,4,5,8,10,12], and (5+7)/2=6 is the median of [1,2,3,5,7,8,10,11].

Using the idea from Lesson 3.2, we can use a pair of heaps to create a data structure which allows fast access to the median. Use the heapq module in python to create both a max-heap and a min-heap. Note that by default, the heapq module in python only creates min-heaps, but if we multiply elements by -1 when we store them, then we can also create max-heaps.


## Question 1.
Write a function `add_to_median_heap(minh, maxh, elem)`. It must accept a min heap, a max heap, and an element to add.


<!-- grade_id: cell-50cb08de70712ee7, points: 0, grade: false, solution: true -->
```python
import heapq

def add_to_median_heap(minh, maxh, elem):
    # YOUR CODE HERe
    if not minh and not maxh:
        heapq.heappush(maxh,-elem)
    else:
        if elem > -maxh[0]:
            heapq.heappush(minh,elem)
        if elem <= -maxh[0]:
            heapq.heappush(maxh,-elem)
    #rebalance two heap lists
    if len(minh) - len(maxh) >= 2:
        switch = heapq.heappop(minh)
        heapq.heappush(maxh,-switch)
    if len(maxh) - len(minh) >= 2:
        switch = heapq.heappop(maxh)
        heapq.heappush(minh,-switch)
#     raise NotImplementedError()
```

```python
lis = [7,5,1,6,3,8,9]
minh = []
maxh = []
import numpy as np
for elem in lis:
    add_to_median_heap(minh, maxh, elem)

print(minh)
print(maxh)
```

Output:
```
[6, 7, 8, 9]
[-5, -1, -3]
```

## Question 2
Write a function `median(minh, maxh)`. It must return the median element.


<!-- grade_id: cell-43ce866dd1dc2d5b, points: 0, grade: false, solution: true -->
```python
def median(minh, maxh):
    if len(minh) == len(maxh):
        median = (minh[0] - maxh[0])/2
    elif len(minh) > len(maxh):
        median = minh[0]
    else:
        median = maxh[0]
    return median
    raise NotImplementedError()
```

```python
median(minh, maxh)
```

Output:
```
6
```

<!-- grade_id: cell-673faa69bf330eda, points: 0, grade: true, solution: false -->
```python
# Please ignore this cell. This cell is for us to implement the tests 
# to see if your code works properly. 
```

## Question 3.

Uncomment and run the testing code given below to test your functions. It should print out numbers ranging from 1 to 50, in that order.

```python
minh = []
maxh = []
for a in range(1,100,2):
    add_to_median_heap(minh, maxh, a)
    print(median(minh, maxh))  
```

Output:
```
-1
2.0
3
4.0
5
6.0
7
8.0
9
10.0
11
12.0
13
14.0
15
16.0
17
18.0
19
20.0
21
22.0
23
24.0
25
26.0
27
28.0
29
30.0
31
32.0
33
34.0
35
36.0
37
38.0
39
40.0
41
42.0
43
44.0
45
46.0
47
48.0
49
50.0
```

## Question 4.
What’s the worst case complexity to build a median heap using `add_to_median_heap`?

<!-- grade_id: cell-c8e33ebdc8aaf328, points: 0, grade: true, solution: true -->
Supposing x1 and x2 be the size of minh and maxh. The worst case occurs as the new element is added in minh, then another element in minh is popped out to be heap pushed in maxh. In this case, the time complexity is lg(x1) + lg(x2) becasue the compexity of heappush is lg(n)

## Question 5.
What’s the worst case complexity of `median`?

<!-- grade_id: cell-c061920faa337890, points: 0, grade: true, solution: true -->
The worst case occurs when the worst case of add_to_median_heap occurs to every element in the input list (reversedly-sroted input can cause this case). Therefore the worst time complexity can be generalizaed as below:
* first element to max heap: constant
* first element in min heap: constant
* second elelment to min heap): O(lg(1))
* third lelment to min: O(lg(2))
* now min heap has 2 element more than max heap(3 vs 1), one element is popped out from minh to be the second element in max heap: ((lg(1))
* third lelment to min: O(lg(2))
* forth lelment to min: O(lg(3))
* now min heap has 2 element more than max heap (4 vs 2 elements), one element is popped out from minh to be the third element in max heap: O((lg(2))
.....
* the process ends when both heap lists have the same n/2 elements. This means running time up until this moment is 
    * cost of minh = constant + O(lg(1)) + O(lg(2))+O(lg(2)) + O(lg(3)) + O(lg(3)) + .. + O(lg(n/2)+O(lg(n/2)) = O(lg(n/2)
    * cost of maxh = constant + O(lg(1)) + O(lg(2)) + O(lg(3)) + .. +O(lg(n/2)) = O(lg(n/2))
* Total cost is O(lg(n/2) + O(lg(n/2) = O(lg(n). Therefore, the worst case running time is O(lg(n)

## Question 6.

How does this way of finding the median compare with the vanilla way of sorting the list and pick the middle element? Use arguments based on efficiency or clearity of the respective algorithms.

<!-- grade_id: cell-1f31afecda9730a7, points: 0, grade: true, solution: true -->
This method is way better than tradition methods in which we have to sort out the list before finding the median. Even if the most optimal sorting algoriths are applied, such as heapsort, their costs are still O(nlgn), compared to O(lg(n)) of the above method.

## [Optional] Question 7.

Is it possible to extend this idea to any percentile? If it is, then write code to do so. If it’s not possible, prove why it is not possible.

<!-- grade_id: cell-4d0c197e6d484338, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

## Part B. Quickselect

Quicksort can be modified to find the $k$-th smallest element in an unordered list. This is known as quickselect. It does this by choosing a partition (as in quicksort). Once the list has been partitioned then we know how many elements lie to the left and to the right of the partition. This allows us to recursively call quickselect on the correct sublist.

## Question 1.

Write a function `qselect(lst, k)`, which takes both a list and an index $k$. The function must then return the $k$-th smallest item in the list.


```python
import timeit
import random

eps = 1e-16

locations = [0.0, 0.5, 1.0 - eps]

# find the median of three values x1, x2, x3.
#Input: a list of 3 number
#Output: the median number of three inputs
def median(x1, x2, x3):
    if (x1 < x2 < x3) or (x3 < x2 < x1):
        return x2
    elif (x1 < x3 < x2) or (x2 < x3 < x1):
        return x3
    else:
        return x1

#define quickshort function
#input: an arbitrary list
#output: sorted list

def qsort(lst,k):
    indices = [(0, len(lst))] #indice is a tuple of 2 elements

    while indices: #while indice is not empty
        (frm, to) = indices.pop() #assign from = 0 and to = len(lst)
        if frm == to: #if from = to skip the following operations
            continue

        # Find the partition:
        N = to - frm #N is the number of element in the list
        inds = [frm + int(N * n) for n in locations] #inds is a list of values following this rule
        values = [lst[ind] for ind in inds] #[a list of elements in the list whose indexes are in inds]
        partition = median(*values) #partition is the median value of the list values

        # Split into lists:
        lower = [a for a in lst[frm:to] if a < partition] #split the list to 2 using the median threshold
        upper = [a for a in lst[frm:to] if a > partition]
        #count the frequence of the median value in the list
        counts = sum([1 for a in lst[frm:to] if a == partition]) 

        ind1 = frm + len(lower) #the ending index of list 1
        ind2 = ind1 + counts #the ending index of the sublist of all median values (in line 48)

        # Push back into correct place:
        lst[frm:ind1] = lower #first part having all value smaller than the median value
        lst[ind1:ind2] = [partition] * counts #the median value
        lst[ind2:to] = upper #the third part of having all values larger the median value

        # Enqueue other locations
#         append indexes of the first and third part of the list in indices
        if k in range(ind1,ind2):
            return lst[ind1]
            break
        elif k < ind1:
            indices.append((frm, ind1)) 
        elif k >=ind2:
            indices.append((ind2, to))
```

<!-- grade_id: cell-fb9c2bc9d5b4bc92, points: 0, grade: false, solution: true -->
```python
def qselect(lst,k):
    if k > len(lst):
        print('k must be not larger than the length of input')

    else:
        return(qsort(lst,k))
#     raise NotImplementedError()
```

```python
qsort(lst1,1000)
```

<!-- grade_id: cell-85a25ed6a1f62a66, points: 0, grade: true, solution: false -->
```python
import random
random.seed(123) # introducing a seed for reproducibility purposes
lst1 = list(range(100))
random.shuffle(lst1)
lst2 = []
for a in range(100):
    lst2.append(qselect(lst1, a))
assert(lst2 == list(range(100)))
```

## Question 2.
Uncomment and run the testing code given below to test your function. It should print out integers from 0 to 99.


```python
import random
random.seed(123) # introducing a seed for reproducibility purposes
lst = list(range(100))
random.shuffle(lst)
for a in range(100):
    print(qselect(lst, a))

```

Output:
```
0
1
2
3
4
5
6
7
8
9
10
11
12
13
14
15
16
17
18
19
20
21
22
23
24
25
26
27
28
29
30
31
32
33
34
35
36
37
38
39
40
41
42
43
44
45
46
47
48
49
50
51
52
53
54
55
56
57
58
59
60
61
62
63
64
65
66
67
68
69
70
71
72
73
74
75
76
77
78
79
80
81
82
83
84
85
86
87
88
89
90
91
92
93
94
95
96
97
98
99
```

## Question 3.

Write down the recurrence relation for your code.

<!-- grade_id: cell-24144607b3ffade3, points: 0, grade: true, solution: true -->
The code above does not technically recall itself, hence not a recursion. But the fact that tuple of ideices (frm:to) are added to lidices list can be seen as a recursion as the code is repeated on different part of the input list. Particularly, After a run, a original list is broken down to 3 parts (numbers smaller than partion, partitions, and numbers larger than partiion). Depending on indcies of those sublists compared to k, the corresponding sublist's indeces would be added to the indeces list. Accordingly, the code would be repeated (recurs) on that sublist

## Question 4.

Solve the recurrence relation for quickselect in the best case.


<!-- grade_id: cell-c125f4742933e820, points: 0, grade: true, solution: true -->
In the best case scenario, the k-th smallest element is the first pivot of the partition (which can be the beggining, the middle and the ending element in the list). In this case, the cost to run qselect is the cost to run qsort() once, which is O(n)

## Question 5.

Solve the recurrence relation for quickselect in the worst case.

<!-- grade_id: cell-852c3a71d8c2d7c1, points: 0, grade: true, solution: true -->
* In the worst case, the k-th smallest element is the pivot in the last partition. To reach this base case, lg(n) partitions must be executed. The reason is that every partition produces balance sublists in size due to the median of three method. So, we can assume each sublist after a partiion has half of the size of the input. Also, in the last partition, the inputsize must be 1. Using those two assumption, the number of partition can be found by solving this equation: 2^x = n => x = lg(n).
* Another assumption in the worse case scenario is the length of the sublist having all partitions is negligible so that the sublist of numbers larger and smaller than the pivot is half of the size of the input.
* As discussed in class 6.2, the cost to run one traditional parition as in the algorithm of qsort() is O(n). In this modefied vesion, the cost should be O(n/2) as the partion is recalled one sublist only. Therefore, wosst case running time is O(n) + O(n/2) + O(n/4) + ... + O(1) = O(2^(lgn + 1) - 1) = O(n)

```python

```